---
## Part 1 (1 points)

Write the following mathematical formula in Python:

\begin{align}
 result = 6a^3 - \frac{8b^2 }{4c} + 11
\end{align}


<!-- grade_id: p1, points: 0, grade: false, solution: true -->
```python
a = 2
b = 3
c = 2

### Store the solution in the variable result
# result = 
### BEGIN SOLUTION
result = 6 * a ** 3 - 8 * b ** 2 / (4 * c) + 11
### END SOLUTION
```

<!-- grade_id: p1-t1, points: 1, grade: true, solution: false -->
```python
assert result == 50
```

---
## Part 2 (1 points)

Floating point pitfalls 
Show that `0.1 + 0.1 + 0.1 == 0.3`

HINT: decimal

<!-- grade_id: p2, points: 1, grade: true, solution: true -->
```python
### BEGIN SOLUTION
from decimal import Decimal

x = Decimal('0.1')
y = Decimal('0.1')
z = Decimal('0.1')

float(x + y + z)
### END SOLUTION
```

Output:
```
0.3
```

---
## Part 3 (1 points)

Filter a list and get only the uppercase elements

<!-- grade_id: get_upper, points: 0, grade: false, solution: true -->
```python
def get_upper(words):
    ### BEGIN SOLUTION
    return [word for word in words if word.isupper()]
    ### END SOLUTION
```

<!-- grade_id: get_upper_test, points: 1, grade: true, solution: false -->
```python
words = ['PYTHON', 'JOHN', 'chEEse', 'hAm', 'DOE', '123']

assert get_upper(words) == ['PYTHON', 'JOHN', 'DOE']
```

---
## Part 4 (1 points)

Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    ### BEGIN SOLUTION
    if n < 1:
        raise ValueError("n must be greater than or equal to 1")
    return [i ** 2 for i in range(1, n + 1)]
    ### END SOLUTION
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
try:
    squares(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    squares(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

---

## Part 5 (1 point)

Using your `squares` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `squares` function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    ### BEGIN SOLUTION
    return sum(squares(n))
    ### END SOLUTION
```

<!-- grade_id: correct_sum_of_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
assert sum_of_squares(1) == 1
assert sum_of_squares(2) == 5
assert sum_of_squares(10) == 385
assert sum_of_squares(11) == 506
```

<!-- grade_id: sum_of_squares_uses_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""
orig_squares = squares
del squares
try:
    sum_of_squares(1)
except NameError:
    pass
else:
    raise AssertionError("sum_of_squares does not use squares")
finally:
    squares = orig_squares
```