Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

---
## Part 1 (1 points)

Fill `____` pieces below to have correct values for `lower_cased`, `stripped` and `stripped_lower_case` variables.

<!-- grade_id: cell-56db8d8fcc3971ee, points: 0, grade: false, solution: true -->
```python
original = ' Python strings are COOL! '
lower_cased = original._____
stripped = ____.strip()
stripped_lower_cased = original._____._____

# DELETE THE RAISE EXCEPTION BELOW
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a060047614e2c6dd, points: 1, grade: true, solution: false -->
```python
assert lower_cased == ' python strings are cool! '
assert stripped == 'Python strings are COOL!'
assert stripped_lower_cased == 'python strings are cool!'
```

---
## Part 2 (1 points)

Given a string of odd length >= 5, return a string composed of the three middle characters of the input.

<!-- grade_id: cell-b0560df0906b50b1, points: 0, grade: false, solution: true -->
```python
def middle_chars(orig):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b2449c576093e63f, points: 1, grade: true, solution: false -->
```python
assert middle_chars('12345') == '234'
assert middle_chars('abc000ghi') == '000'
```

---
## Part 2 (1 points)

Find the first appearance of the substring 'not' and 'poor' from a given string, if 'not' follows the 'poor', replace the **whole** 'not'...'poor' substring with 'good'. Return the resulting string

<!-- grade_id: cell-79275f6d6146ff60, points: 0, grade: false, solution: true -->
```python
def replace_str(orig):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e7a4e71debcf7ca0, points: 1, grade: true, solution: false -->
```python
assert replace_str('The lyrics are not that poor!') == 'The lyrics are good!'
assert replace_str('The lyrics are poor!') == 'The lyrics are poor!'
```

---
## Part 3 (1 points)
Find if all chars in s1 are in s2.

<!-- grade_id: cell-423fa6f68129ad12, points: 0, grade: false, solution: true -->
```python
def s1_in_s2(s1, s2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f20d039aee4b3a2e, points: 1, grade: true, solution: false -->
```python
assert s1_in_s2('12345', '111222333444555') == True
assert s1_in_s2('12345', '222333444555') == False
```

---
## Part 4 (1 point) 
---
**one-liner bonus** (1 point)

Write a Python program that accepts a comma separated sequence of words and prints the 
unique words in sorted form (alphanumerically).



<!-- grade_id: cell-290e46d5aadbb1a3, points: 0, grade: false, solution: true -->
```python
def sort_words(string):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8877572ac2227f28, points: 1, grade: true, solution: false -->
```python
assert sort_words('red, black, pink, green, black, green, pink, red') == 'black, green, pink, red'
```

---
# Part 5 (1 points)
Write a Python taking a cube edge size and the number of decimal precision we want and returns the volume of the cube in cm^3.
Look online for unicode symbols.

<!-- grade_id: cell-80a6bd7b9b4b5004, points: 0, grade: false, solution: true -->
```python
def volume_of_cube(edge, decimals):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-14d588762975409c, points: 1, grade: true, solution: false -->
```python
assert volume_of_cube(2, 2) == 'The volume of the cube is 8.00cm³'
assert volume_of_cube(4.5, 3) == 'The volume of the cube is 91.125cm³'
```

---
# Part 6 (2 points)
Count repeated characters of a string in a dictionary and order the resulting array of touples by number of occurances from higher to lower and in the case of equality order alphabetically.

ex:
```
in: repeated('0001112223334455')
out: [('0', 3), ('1', 3), ('2', 3), ('3', 3), ('4', 2), ('5', 2)]
```

<!-- grade_id: cell-2a796b97b4e8ad5b, points: 0, grade: false, solution: true -->
```python
def repeated(string):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fd35022c43c67491, points: 2, grade: true, solution: false -->
```python
assert repeated('asdasddsada') == [('a', 4), ('d', 4), ('s', 3)]
assert repeated('0001112223334455') == [('0', 3), ('1', 3), ('2', 3), ('3', 3), ('4', 2), ('5', 2)]
```