---
## Part 1 (1 points)
Return a tuple containing the sum, mean average and length of a list given as input.

<!-- grade_id: cell-1dd9ebecc6447752, points: 0, grade: false, solution: true -->
```python
def average(l1):
    ### BEGIN_SOLUTION
    s = sum(l1)
    l = len(l1)
    return s, s / l, l
    ### END_SOLUTION
```

<!-- grade_id: cell-9486a334677c53f5, points: 1, grade: true, solution: false -->
```python
assert average([1, 2, 3]) == (6, 2.0, 3)
```

---
## Part 2 (1 points)
Create a third list from l1 and l2 that contains the even indexed elements from the first appended to the odd indexed elements from the second

<!-- grade_id: cell-75325ad0d6cc3321, points: 0, grade: false, solution: true -->
```python
def l1_add_l2(l1, l2):
    ### BEGIN_SOLUTION
    return [*l1[0::2], *l2[1::2]]    
    ### END_SOLUTION
```

<!-- grade_id: cell-cf373b20f2e2dc1d, points: 1, grade: true, solution: false -->
```python
assert l1_add_l2([1, 2, 3, 4, 5],[1, 2, 3, 4, 5]) == [1, 3, 5, 2, 4]
```

---
## Part 3 (1 points)
Remove an element at position i and put it at position j.

<!-- grade_id: cell-ffc58f4512d2d0d5, points: 0, grade: false, solution: true -->
```python
def pop_and_add(s, i, j):
    ### BEGIN_SOLUTION
    item = s.pop(i)
    s.insert(j, item)
    return s
    ### END_SOLUTION
```

<!-- grade_id: cell-2ffa9f0053765139, points: 1, grade: true, solution: false -->
```python
assert pop_and_add([1, 2, 3], 2, 0) == [3, 1, 2]
assert pop_and_add([1, 2, 3, 4, 5], 1, 4) == [1, 3, 4, 5, 2]
```

---
## Part 4 (1 points)
Given two sets s1 and s2, return a set containing elements from s1 or s2 but not both

<!-- grade_id: cell-5b3ceb6ea2878f2d, points: 0, grade: false, solution: true -->
```python
def set_ops(s1, s2):
    ### BEGIN SOLUTION
    return s1 ^ s2
    ### END SOLUTION
```

<!-- grade_id: cell-fdbb5305e43cbfdb, points: 1, grade: true, solution: false -->
```python
assert set_ops({1,2,3}, {2,3,4}) == {1,4}
```

---
## Part 5 (1.5 points)
Given a list l1, slice it into equally sized n chunks.
(The length of the list is always divisible by n)

<!-- grade_id: cell-0049e43676eff8d0, points: 0, grade: false, solution: true -->
```python
def nslice(l1, n):
    ### BEGIN SOLUTION
    size = len(l1) // n    
    return [l1[i:i + size] for i in range(0, len(l1), size)]
    ### END SOLUTION
```

<!-- grade_id: cell-6de93cb32b4e080b, points: 1.5, grade: true, solution: false -->
```python
assert nslice([1, 2, 3, 4, 5, 6], 3) == [[1, 2], [3, 4], [5, 6]]
assert nslice([1, 2, 3], 3) == [[1], [2], [3]]
```

---
## Part 6 (1.5 points)
Given multiple lists as arguments create a function that returns the cartesian product.
ex:
```
in: [1, 0], [2, 3]
out: [[1, 2], [1, 3], [0, 2], [0, 3]]
```


<!-- grade_id: cell-8a122ea79bb7e70a, points: 0, grade: false, solution: true -->
```python
def cart_prod(*lists):
    ### BEGIN SOLUTION
    result = [[]]
    for l in lists:
        result = [x + [y] for x in result for y in l]
        # print (result)
        
    return result
    ### END SOLUTION
```

<!-- grade_id: cell-49c72a78beb51f34, points: 1.5, grade: true, solution: false -->
```python
assert cart_prod([1, 0], [2, 3]) == [[1, 2], [1, 3], [0, 2], [0, 3]]
assert cart_prod([1, 0], [2, 3], [4, 5]) == [[1, 2, 4], [1, 2, 5], [1, 3, 4], [1, 3, 5], 
                                             [0, 2, 4], [0, 2, 5], [0, 3, 4], [0, 3, 5]]
```

---
## Part 7 (1.5 points)
Create all possible combinations of letters given a dictionary which values are characters.  
**HINT**: use something implemented before  

ex:
```
in: {1 : ['a','b'], 2 : ['c','d']}
out: ['ac', 'ad', 'bc', 'cd']
```

<!-- grade_id: cell-55b26f0accdb7d23, points: 0, grade: false, solution: true -->
```python
def all_chars(d1):
    ### BEGIN SOLUTION
    return [''.join(x) for x in cart_prod(*d1.values())]
    ### END SOLUTION
```

<!-- grade_id: cell-484620fc342d2353, points: 1.5, grade: true, solution: false -->
```python
in0 = {1 : ['a','b'], 2 : ['c','d']}
out0 = ['ac', 'ad', 'bc', 'bd']
assert sorted(all_chars(in0)) == out0

in1 = {1 : ['a','b'], 2 : ['c','d'], 3 : ['e']}
out1 = ['ace', 'ade', 'bce', 'bde']
assert sorted(all_chars(in1)) == out1
```

---
## Part 8 (1 points)
Get the first 3 biggest values from a dictionary and return a list of the keys. 

ex:
```
in: {1 : 12, 2 : 24, 3: 30, 4: -1, 5: 22}
out: [3, 2, 5]
```

<!-- grade_id: cell-a972d28f05bcfa20, points: 0, grade: false, solution: true -->
```python
def biggest_3(d1):
    ### BEGIN SOLUTION
    return [x for x, _ in sorted(d1.items(), key=lambda x: x[1], reverse=True)[:3]]
    ### END SOLUTION
```

<!-- grade_id: cell-255d924b8ffff5b8, points: 1, grade: true, solution: false -->
```python
assert biggest_3({1 : 12, 2 : 24, 3: 30, 4: -1, 5: 22}) == [3, 2, 5]
```

---
## Part 9 (1 points)
Given two dictionaries d1 and d2, return a dictionary 
that contains all the keys in both the dictionaries that share the same value

ex:
```
in: d1 = {1 : 12, 2 : 24, 3: 30, 4: -1, 5: 22}
    d2 = {1 : 12, 2 : 24, 3: 30, 4: -2, 6: 22}
out: {1 : 12, 2 : 24, 3: 30}
```

<!-- grade_id: cell-6ae3405af3acad74, points: 0, grade: false, solution: true -->
```python
def dict_intersect(d1, d2):
    ### BEGIN SOLUTION
    return {k : d1[k] for k in set(d1.keys()) & set(d2.keys()) if d1[k] == d2[k]}
    ### END SOLUTION

```

<!-- grade_id: cell-7d425ee7d03e52c1, points: 1, grade: true, solution: false -->
```python
assert dict_intersect({1 : 12, 2 : 24, 3: 30, 4: -1, 5: 22},
                      {1 : 12, 2 : 24, 3: 30, 5: -1, 4: 22}) == {1 : 12, 2 : 24, 3: 30}
```

---
## Part 10 (1 points)
Sum all values of key `ammount` for every distinct key `item`
and return the sum as a dictionary, given a list of dictionaries.

ex:
```
in: [{'item': 'item1', 'amount': 400}, 
     {'item': 'item2', 'amount': 300}, 
     {'item': 'item1', 'amount': 750}]
out: {'item1': 1150, 'item2': 300}
```

<!-- grade_id: cell-c777623ed911c552, points: 0, grade: false, solution: true -->
```python
def dict_collect(dictL):
    ### BEGIN SOLUTION
    result = {}
    
    for d in dictL:
        result[d['item']] = result.get(d['item'], 0) + d['amount']
        
    return result
    ### END SOLUTION

```

<!-- grade_id: cell-a837cb5bb8de09f4, points: 1, grade: true, solution: false -->
```python
assert dict_collect([{'item': 'item1', 'amount': 400}, 
                     {'item': 'item2', 'amount': 300}, 
                     {'item': 'item1', 'amount': 750}]) == {'item2': 300, 'item1': 1150}
```