Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

---
### Part 1 (1 points)

Create a function that checks if a number is perfect or not.

According to Wikipedia : In number theory, a perfect number is a positive integer that is equal to the sum of its proper positive divisors, that is, the sum of its positive divisors excluding the number itself (also known as its aliquot sum). Equivalently, a perfect number is a number that is half the sum of all of its positive divisors (including itself).

The first perfect number is 6, because 1, 2, and 3 are its proper positive divisors, and `1 + 2 + 3 = 6`

<!-- grade_id: cell-c61814b57c00d2a2, points: 0, grade: false, solution: true -->
```python
def perfect(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3ab9384409fe521b, points: 1, grade: true, solution: false -->
```python
assert perfect(6) and perfect(28) and not perfect(12)
```

---
### Part 2 (1 points)

Create a function that returns a function that appends a given string to the end and front of a string.

ex:
```
bold = appender('<bold>')
bold('hi') -> <bold>hi<bold>
```

<!-- grade_id: cell-25e1f113f186d8ce, points: 0, grade: false, solution: true -->
```python
def appender(st):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8593888991e19abc, points: 1, grade: true, solution: false -->
```python
bold = appender('<bold>')
assert bold('hi') == '<bold>hi<bold>'

bold = appender(' *** ')
assert bold('title') == ' *** title *** '
```

---
### Part 3 (1 points)

Create a function that sums all list args and multiplies all dictionary args and returns a tuple

ex:
```
bold = appender('<bold>')
bold('hi') -> <bold>hi<bold>
```

<!-- grade_id: cell-35c532cb23f073b1, points: 0, grade: false, solution: true -->
```python
# Write the function definition -> def args_adder(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f6855dd7b4d17dc8, points: 1, grade: true, solution: false -->
```python
assert args_adder(*range(0, 11), asd=1, bce=2, cdf=3) == (55, 6)
assert args_adder(asd=1, bce=2, cdf=3) == (0, 6)
```

---
### Part 4 (1 points)

Write this simple division function and handle all possible exceptions.  
Return a string representing the exception if an exception occurs.
```
zero division -> one item is zero
bad type -> invalid type
```

<!-- grade_id: cell-4673f51220303cdf, points: 0, grade: false, solution: true -->
```python
def simple_div(x, y):
    result = None # x / y
    # YOUR CODE HERE
    raise NotImplementedError()
    return result
```

<!-- grade_id: cell-677478a0e8835be3, points: 1, grade: true, solution: false -->
```python
assert simple_div(1, 3) == 1 / 3
assert simple_div(1, 0) == 'one item is zero'
assert simple_div('asd', 0) == 'invalid type'
```

---
### Part 5 (1 points)

From a lists of inner lists of unknown depth flatten to a single list.

```
in: [[3, [1, 2], [[[6], 5], 4, 0], 7, [[8]], [9, 10]]]
out: [3, 1, 2, 6, 5, 4, 0, 7, 8, 9, 10]
```

<!-- grade_id: cell-ffd9101c90ff63a0, points: 0, grade: false, solution: true -->
```python
def flatten(l1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-666621f38ef428a5, points: 1, grade: true, solution: false -->
```python
assert flatten([[3, [1, 2], [[[6], 5], 4, 0], 7, [[8]], [9, 10]]]) == [3, 1, 2, 6, 5, 4, 0, 7, 8, 9, 10]
assert flatten([[], [], [], 1]) == [1]
```

---
## Part 6 (1.5 points)

Return the n-th fibonnaci number
```
idx: 1, 2, 3, 4, 5, 6
fib: 1, 1, 2, 3, 5, 8...
```

<!-- grade_id: cell-68100df7fb21de0b, points: 0, grade: false, solution: true -->
```python
def fib(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b46d7eee878c29e2, points: 1.5, grade: true, solution: false -->
```python
assert fib(1) == 1 and fib(2) == 1 and fib(12) == 144
'''If the second assert fails think about why :D'''
assert fib(72) == 498454011879264
```